# Setup

CS1302 Introduction to Computer Programming
___

## Course Materials

### How to access the programming environment?

The course's programming environment is conveniently accessible via a JupyterHub server, allowing remote access without the need to install any special software, thanks to <wiki:Project_Jupyter>. This means you can easily write and run programs on your mobile devices using just a web browser.

Open the url of the JupyterHub server given in the course homepage. The first time you access the url, it opens a login page as shown in [](#fig:login-box):
  1. Enter your [EID](https://www.cityu.edu.hk/esu/eid.htm) and Password in the fields `Username` and `Password` respectively.
  1. Click the `Sign in` button.

:::{figure} images/login.dio.svg
:name: fig:login-box
:alt: Login box
:align: left

Login box
:::

To access the Jupyter environment, simply log into the JupyterHub server and select the `Default` option from the list of available Jupyter servers. Click `Start` to begin your session. If the server is spawned successfully, the JupyterLab interface should appear. Each student will have their own Jupyter server, providing individual computer resources such as CPU and memory.

::::{note} Server options

- The `Default` option runs on the Ubuntu operating system and offers a wide range of features.
  See the [dockerfile](https://github.com/dive4dec/cs1302jupyter/blob/main/cs1302nb/Dockerfile) for details.
- Alternatively, the `Alpine` option may be faster, but it has fewer features. It runs on a lightweight Linux distribution called Alpine Linux.
  See the [dockerfile](https://github.com/dive4dec/cs1302jupyter/blob/main/cs1302anb/Dockerfile) for details.
- You can switch between the `Default` and other server options at any time by restarting your server and selecting the desired option. Restarting the server is simple and can be done through the hub control panel, as described in the note below.

:::{tip} Troubleshooting[^1]
:class: dropdown

In case the jupyterlab interface fails to show:

- Refresh the page or restart your browser.
- Restart your server using the `Hub Control Panel` under the `File` and click `Stop My Server` and then `Start My Server`.

:::

::::

[^1]:  click me to expand

### How to access the course materials?

Course materials are written in the form of Jupyter Notebooks. Jupyter Notebooks provide an interactive literate programming experience, which is a significant benefit for learners:

- Jupyter Notebooks allow programs to be mixed with other contents, including figures, videos, and formatted textual explanations.
- The notebooks can also run in a Jupyter server, which enables users to write and run their programs while adding formatted notes to it.

The recommended way is to follow the links to the notebooks on the course homepage as shown in [](#fig:notebook-links).

:::{figure} images/canvas.dio.svg
:name: fig:notebook-links
:alt: Links to lecture/lab notebooks
:align: left

Links to lecture/lab notebooks on course homepage
:::

The notebooks reside in a [GitHub repository](https://github.com/dive4dec/cs1302_23a/blob/main/README.md). The first time you access the notebook link, the following steps occur:

- Your Jupyter server is opened.
- The repository is cloned to the server as a local repository.
- The desired notebook is opened in the JupyterLab interface.

Afterwards, accessing a notebook link again will

- [git-pull](https://git-scm.com/docs/git-pull)s new or existing files in the repository and 
- merge them with the notebooks stored under the course folder `cs1302_xxx` in your home directory without overwriting your changes.

::::{important}

Manually fetching the notebooks through other means may not create the local repository required for git version control. This can result in new notebooks or changes to existing notebooks not being pulled properly to your server. To ensure that you have the latest versions of the notebooks with proper version control, use the provided notebook links.

:::{seealso}

If you are interested in how the link works, see [`nbgitpuller`](https://jupyterhub.github.io/nbgitpuller/index.html). A git-pull link looks as follows:

```bash
https://{{server_url}}/git-pull        # send request to jupyter server
repo={{git_repo}}&                     # pull from the git repository
urlpath=lab/tree/{{path_to_notebook}}  # follow the path to open the notebook
```

:::

::::

## Lab Assignments

### How to complete a lab assignment?

The notebooks can be edited in JupyterLab to include your notes and answers for submission. If this is your first time using JupyterLab, take a look at the official video tutorial:

:::{iframe} https://www.youtube.com/embed/A5YyoCKxEOU
:width: 100%
:align: left

Official video tutorial on Jupyter
:::

The `Help` menu gives more detailed references. For instance:

- Checkout the menu item `Show Keyboard Shortcuts` and try some of the shortcuts to see their effect. `Jupyter Reference` opens the [user guide](https://jupyterlab.readthedocs.io/en/latest/user/interface.html) in a new tab.
- Try the [MyST Markdown](https://mystmd.org/guide/typography) syntax to format your notes.

In learning a new programming language, the first program to write is often the ["Hello, World!"](https://en.wikipedia.org/wiki/%22Hello,_World!%22_program) program, which says Hello to the world. As your first lab exercise, you will write such a program in python.

:::::{exercise} Hello World
:label: ex:Hello-World

Complete the program in the solution cell to print the message `Hello, World!`. 

::::{hint}
:class: dropdown

One possible solution is:

:::{code-block} python
:name: code:Hello-World
:caption: Python, "Hello, World!"

def say_hello():
    print("Hello, World!")
    
say_hello()
:::

:::{caution}
You must use a tab or 4 spaces to indent the second line of code `print(...)`.
:::

::::

:::::

The following code cell is the solution. Since you are not expected to know python before, you can simply expand the hint above and copy the solution to the solution cell instead.

<!-- grade_id: helloworld, points: 0, grade: false, solution: true -->
```python
def say_hello():
    # YOUR CODE HERE
    raise NotImplementedError()
    
say_hello()
```

:::{important}

It's important to follow certain guidelines when writing your answers or code in the notebooks:

- Only provide your answers in the cells that are designated for this purpose, which are usually labeled with `YOUR CODE HERE` or `YOUR ANSWER HERE`.
- For coding exercises, be sure to remove the line `raise NotImplementedError()` from the cell before submitting your work.
- Do not clone or duplicate the answer cells, as this can cause issues with version control and grading.

:::

To test your program:

- Run your program by selecting the solution cell and press `Shift`+`Enter`.
- Then, run the following visible test to check whether your program prints the correct message.

<!-- grade_id: test-helloworld, points: 1, grade: true, solution: false -->
```python
# Run this test cell right after running your "Hello, World!" program.
import io
import sys

old_stdout, sys.stdout = sys.stdout, io.StringIO()
say_hello()
printed = sys.stdout.getvalue()
sys.stdout = old_stdout
assert printed == "Hello, World!\n"
```

The test returns an assertion error if your program does not print the correct message.

```{tip}

- You can repeatedly modify your solution and run the test cell until your solution passes the test. There is no mark penalty in failing the test before submission. 
- You should try to understand the error messages from a failed test.
- To assess your solution thoroughly, we will run new tests hidden from you after you have submitted your notebook. Therefore, *you should ensure your solution works in general rather than just the visible tests*.
- If you open the same notebook multiple times in different browser windows, be careful in making changes in different windows as you may overwrite your own changes.
- If your notebook fails to run any code, the kernel might have died. You can restart the kernel with `Kernel`$\to$ `Restart`. If restarting fails, check your code cells to see if there is anything that breaks the kernel.

```

### How to submit a notebook

:::{important} Grading policies

- Late submissions will not be accepted without valid justifications.
- You can submit your assignment repeatedly before the deadline without penalty, so it is recommended to submit your assignment in advance. Please note that you are responsible for any issues related to failed submissions due to network outages that occur too close to the deadline.
- It is also important to make a record or backup of your submission that includes the submission timestamp. Double check to ensure that you have submitted the correct lab assignment, since multiple lab assignments may be open for submission at the same time.

:::

You normally have at least 5 days to work on the lab after your lab session. 

- You can check the due dates of all the labs from the course homepage. 
- You may seek help from us or your classmates. However, you must write your own solution and indicate who your collaborators are by including their EIDs in a code cell:  
    ```Python
    COLLABORATORS = ["xfwong2", "mklee3"]
    ```

:::{attention}

- Although the first Lab (`Lab0`) does not count towards your final grade, you should still submit it, to get familiar with the procedure.
- Lecture notebooks under the subfolders `Lecture*\` need NOT be submitted. You only need to submit the lab notebooks under `Lab*\`. 

:::

[Nbgrader](https://nbgrader.readthedocs.io/en/stable/index.html) is the tool used for grading notebook assignments. It is well-integrated into the JupyterLab interface, allowing students to submit their notebooks directly through the JupyterHub server. Submitted notebooks can be both auto-graded with pre-defined test cases and manually graded with custom feedback. After grading is complete, you can check your scores and access the feedback using the same interface.

Before you submit, make sure everything runs as expected:

1. **Git-pull the notebooks** (optional): Follow any one of the link on the course homepage to a notebook, which will git-pull any updates/corrections to (all) your notebooks.
1. **Save the notebook**: Unsaved changes are not exported, even though they are in the memory.
1. **Restart the kernel**: `Kernel`$\to$ `Restart Kernel...` to have a clean state before running visible tests.
1. **run all cells**: `Run`$\to$ `Run All Cells` to double check the results of the visible tests are as expected.

To submit your notebook:

1. Select the menu item `Nbgrader`$\to$`Assignment List`.
1. Expand the Lab folder and click the `validate` button next to the notebook(s) to check if all the visible tests pass.
1. Click `Submit` to submit your notebook.

::::{caution}

Your submission may not be graded under the following circumstances:

- If the notebook files have been renamed, for example, `Setup.ipynb` being changed or copied to `Setup (Copy).ipynb`.
- If an HTML file exists with the same name as a notebook, for example, `Setup.html`.
- If the file size exceeds `100MB`.
- If the code takes too long to run or requires an excessive amount of memory.

It is essential to ensure that your submission meets these guidelines to avoid any issues with grading.

:::{tip} Troubleshooting
:class: dropdown

If you believe your notebooks are corrupted, 
1. download/backup your existing notebooks,
1. remove them from the `Lab*/` folder,
1. click the git-pull links from the course homepage to re-pull the notebooks, and
1. copy your solutions to the new notebooks.

You may also run the course notebooks outside the jupyterhub server and locally on your computer. For more details, see the course homepage.

:::

::::