# Improved Quadratic Equation Solver

CS1302 Introduction to Computer Programming
___

```python
import math
import numpy as np
from ipywidgets import interact
```

In this notebook, we will improve the quadratic equation solver in the previous lab using conditional executions.

## Discriminant

Recall that a quadratic equation is

$$
ax^2+bx+c=0
$$ (eq:quadratic)

where $a$, $b$, and $c$ are real-valued coefficients, and $x$ is the unknown variable.

::::{prf:definition} Discriminant


The disciminant of a quadratic equation is defined as

$$
\Delta := b^2 - 4ac.
$$ (eq:discriminant)

::::

The discriminant $\Delta$ discriminates the roots 

$$\frac{-b\pm \sqrt{\Delta}}{2a}$$ (roots)

of a quadratic equation:

::::{prf:proposition}

The roots of a quadratic equation {eq}`eq:quadratic` are both equal to 

$$-\frac{b}{2a}$$ (eq:repeated_root)

if and only if the discriminant is zero, i.e., $\Delta=0$.

::::

For example, if $(a,b,c)=(1,-2,1)$, then 

$$\Delta = (-2)^2 - 4(1\cdot 1) = 0$$

and so the repeated root is

$$
- \frac{b}{2a} = \frac{2}2 = 1.
$$

$(x-1)^2 = x^2 - 2x + 1$ as expected.

::::{exercise}  
:label: zero-determinant

Complete the following function by assigning `roots` only one root when the discriminant is zero. E.g., if `(a, b, c) == (1, -2, 1)`, then `roots` should be assigned the value `1.0` instead of `1.0, 1.0`.

:::{hint}
:class: dropdown

You may use the [`if` statement](https://docs.python.org/3/reference/compound_stmts.html#if) as follows:

```python
def get_roots(a, b, c):
    d = b**2 - 4 * a * c    # discriminant
    if math.isclose(d, 0):
        roots = ...  # repeated root
    else:
        d **= 0.5
        roots = ...
    return root
```

Replace `...` in the above code cell by your code. If you have a better implementation, you need *not* follow the solution template. E.g., you can write your solution entirely using boolean operations without any `if` statement.
:::

::::

<!-- grade_id: zero_determininant, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    d = b**2 - 4 * a * c    # discriminant
    if math.isclose(d, 0):
    # YOUR CODE HERE
    raise NotImplementedError()
    return roots
```

<!-- grade_id: test-zero_determinant, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
assert np.isclose(get_roots(1, 0, 0), 0.0).all()

print(get_roots(1,1,0))
```

::::{exercise}
:label: ex:isclose

Why use `math.isclose(d, 0)` instead of `d == 0`?

::::

<!-- grade_id: isclose, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Linear equation

::::{exercise}
:label: ex:zero-a

Give the name of the error that the formula {eq}`roots` implemented in python will raise when $a=0$. Explain whether it is a runtime error or not.

::::

<!-- grade_id: zero-a, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Nevertheless, the quadratic equation remains valid:

$$
bx + c=0,
$$

the root of which is $x = -\frac{c}b$.

::::{exercise}
:label: linear

Improve the function `get_roots` to return the root $-\frac{c}{b}$ when $a=0$.

:::{hint}
:class: dropdown

Solution template:

```python
def get_roots(a, b, c):
    d = b**2 - 4 * a * c    # discriminant
    if ...:
        roots = ...
    elif math.isclose(d, 0):
        roots = ...  # repeated root
    else:
        d **= 0.5
        roots = ...
    return roots
```

::::

<!-- grade_id: linear, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    d = b**2 - 4 * a * c
    # YOUR CODE HERE
    raise NotImplementedError()
    return roots
```

<!-- grade_id: test-linear, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
assert np.isclose(get_roots(1, 0, 0), 0.0).all()
assert np.isclose(get_roots(1, -2, 1), 1.0).all()
assert np.isclose(get_roots(0, -2, 1), 0.5).all()
```

<!-- grade_id: hidden_test-linear, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

## Degenerate cases

What if $a=b=0$? In this case, the equation becomes

$$
c = 0
$$
which is always satisfied if $c=0$ but never satisfied if $c\neq 0$.

::::{exercise}
:label: degenerate

Improve the function `get_roots` to return root(s) under all the following cases:
- If $a=0$ and $b\neq 0$, assign `roots` to the single root $-\frac{c}{b}$. 
- If $a=b=0$ and $c\neq 0$, assign `roots` to `None`.
- If $a=b=c=0$, there are infinitely many roots. Assign to `roots` the tuple `-float('inf'), float('inf')`.  
    Note that `float('inf')` converts the string `'inf'` to a floating point value that represents $\infty$.
    
:::{caution}
- `None` a *python keyword* that refers to a special object. You should *NOT* regard it as a string or quote it. (~~`"None"`~~).
- `float('inf')` converts the string `'inf'` to a floating point number no smaller than any other floating point numbers. You should *NOT* write `inf`, which needs not refer to `float('inf')`.
:::

::::

<!-- grade_id: degenerate, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    d = b**2 - 4 * a * c
    # YOUR CODE HERE
    raise NotImplementedError()
    return roots
```

<!-- grade_id: test-degenerate, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
assert np.isclose(get_roots(1, 0, 0), 0.0).all()
assert np.isclose(get_roots(1, -2, 1), 1.0).all()
assert np.isclose(get_roots(0, -2, 1), 0.5).all()
assert np.isclose(get_roots(0, 0, 0), (-float('inf'), float('inf'))).all()
assert get_roots(0, 0, 1) is None
```

<!-- grade_id: test_test-degenerate, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

After you have complete the exercises, you can run your robust solver below:

```python
# quadratic equations solver
@interact(a=(-10, 10, 1), b=(-10, 10, 1), c=(-10, 10, 1))
def quadratic_equation_solver(a=1, b=2, c=1):
    print("Root(s):", get_roots(a, b, c))
```