# Calculators

CS1302 Introduction to Computer Programming
___

Run the following to load additional tools required for this lab.  
In particular, the `math` library provides many useful mathematical functions and constants.

```python
import math
from math import cos, exp, log, pi, sin, tan

import jsxgraphs
import matplotlib.pyplot as plt
import numpy as np
import sympy as sp
from ipywidgets import interact

%matplotlib widget
```

The following code is a Python one-liner that creates a calculator.  
Evaluate the cell with `Ctrl+Enter`:

```python
print(eval(input()))
```

Try some calculations below using this calculator:
1. $2^3$ by entering `2**3`;
1. $\frac23$ by entering `2/3`;
1. $\left\lceil\frac32\right\rceil$ by entering `3//2`;
1. $3\mod 2$ by entering `3%2`;
1. $\sqrt{2}$ by entering `2**0.5`; and
1. $\sin(\pi/6)$ by entering `sin(pi/6)`;

For this lab, you will create different calculators. We will first show you a demo. Then, it will be your turn to create the calculators.

## Hypotenuse Calculator

::::{prf:proposition} Pythagoras Theorem
:label: pro:pythagorus

By the Pythagoras theorem, given a right-angled triangle,

![Right-angled triangle](images/pythagoras.dio.svg)

the length of the hypotenuse is

$$
c = \sqrt{a^2 + b^2}
$$ (eq:hypotenuse)

where $a$ and $b$ are the lengths of the other sides of the triangle.

::::

You can verify the theorem using the JSXGraph app below:

```python
jsxgraphs.pythagorean1
```

::::{seealso}

Click the button `scratch` to show and edit the Javascript code. Other mathematical illustrations created using `divewidgets` can be found [here](https://dive4dec.github.io/lab/?path=divemath%2Fdivemath.ipynb) and was presented in [JSXGraph conference](https://jsxgraph.uni-bayreuth.de/conf2022/program/).

::::

The following is an interactive graphical proof:

```python
jsxgraphs.pythagorean2
```

Another interactive proof is as follows:

```python
jsxgraphs.pythagorean3
```

We can define the following function to calculate the length `c` of the hypotenuse when given the lengths `a` and `b` of the other sides: 

:::{code-block} python
:name: code:length_of_hypotenuse
:caption: A function that computes the length of hypotenuse

def length_of_hypotenuse(a, b):
    c = (a**2 + b**2) ** (0.5)  # Pythagoras
    return c
:::

:::{note}

- Equation [](#eq:hypotenuse) in [](#pro:pythagorus) is written as a python expression in [](#code:length_of_hypotenuse) using the exponentiation operator `**`.
- The variable `c` is assigned to the value of the expression using the assignment operator `=`.

:::

::::{exercise} length of hypotenuse

Complete the function below to return the length `c` of the hypotenuse given the lengths `a` and `b`.

:::{hint}
:class: dropdown

See [](#code:length_of_hypotenuse).

:::


:::{note}

- You should still complete this exercise to get the credit even though the answer is obvious. Instead of copy-and-paste the answer, type it yourself.
- Indentation affects the execution of python code. In particular, the assignment statement must be indented to indicate that it is part of the *body* of the function.

:::

::::

<!-- grade_id: length_of_hypotenus, points: 0, grade: false, solution: true -->
```python
def length_of_hypotenuse(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
    return c
```

You can check your code against a few cases listed in the test cell below.

<!-- grade_id: test-length_of_hypotenus, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(length_of_hypotenuse(0, 0), 0)
assert np.isclose(length_of_hypotenuse(3, 4), 5)
assert np.isclose(length_of_hypotenuse(4, 7), 8.06225774829855)
```

::::{tip} If you are curious about the hidden test...
:class: dropdown

The hidden test will look like the following but with a "truely" random random `_seed_`:

```python
rng = np.random.default_rng(_seed_)
a, b = rng.random(2)

assert np.isclose(length_of_hypotenuse(a, b), (a**2 + b**2) ** (0.5))
```
::::

<!-- grade_id: htest-length_of_hypotenus, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

We will use `ipywidgets` to let user interact with the calculator more easily as illustrated in [](#fig:hypotenuse-calculator):

- After running the cell, move the sliders to change the values of `a` and `b`. 
- Observer that the value of `c` is updated immediately.

::::{figure} images/hypotenuse-calculator.gif
:name: fig:hypotenuse-calculator
:alt: The hypotenuse calculator
:align: center

Illustration of the hypotenuse calculator

::::

```python
# hypotenuse calculator
@interact(a=(0, 10, 1), b=(0, 10, 1))
def calculate_hypotenuse(a=3, b=4):
    print('c: {:.2f}'.format(length_of_hypotenuse(a, b)))
```

::::{note}

The hypotenuse is printed up to 2 decimal places using the format specification `{:.2f}`.

::::

## Quadratic Equation

### Graphical Calculator for Parabola

::::{figure} images/plot_parabola.gif
:name: fig:parabola
:alt: The parabola calculator
:align: center

Illustration of the parabola calculator

::::

::::{prf:definition} Parabola
:label: def:parabola

The collection of points $(x,y)$ satisfying the following equation forms a *parabola*:

$$
y=ax^2+bx+c
$$ (eq:parabola)

where $a$, $b$, and $c$ are real numbers called the *coefficients*.

::::

The following plots the parabola with difference choices of coefficients.

```python
jsxgraphs.parabola
```

::::{exercise}
:label: ex:parabola

Given the variables `x`, `a`, `b`, and `c` store the $x$-coordinate and the coefficients $a$, $b$, and $c$ respectively, assign `y` the corresponding $y$-coordinate of the parabola [](#eq:parabola).

::::

<!-- grade_id: get_y, points: 0, grade: false, solution: true -->
```python
def get_y(x, a, b, c):
    # YOUR CODE HERE
    raise NotImplementedError()
    return y
```

To test your code:

<!-- grade_id: test-get_y, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(get_y(0, 0, 0, 0), 0)
assert np.isclose(get_y(0, 1, 2, 1), 1)
assert np.isclose(get_y(0, 2, 1, 2), 2)
assert np.isclose(get_y(1.2, 2, 3, 4), 10.48)
assert np.isclose(get_y(2, 3.3, 4, 5), 26.2)
assert np.isclose(get_y(3, 4.4, 5, 6), 60.6)
```

<!-- grade_id: htest-get_y, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

To run the graphical calculator illustrate in [](#fig:parabola):

```python
# graphical calculator for parabola
fig, ax = plt.subplots()
xmin, xmax, ymin, ymax, resolution = -10, 10, -10, 10, 50
x = np.linspace(xmin, xmax, resolution)
ax.set_title(r"$y=ax^2+bx+c$")
ax.set_xlabel(r"$x$")
ax.set_ylabel(r"$y$")
ax.set_xlim([xmin, xmax])
ax.set_ylim([ymin, ymax])
ax.grid()
(p,) = ax.plot(x, get_y(x, 0, 0, 0))


@interact(a=(-10, 10, 1), b=(-10, 10, 1), c=(-10, 10, 1))
def plot_parabola(a, b, c):
    p.set_ydata(get_y(x, a, b, c))
```

### Quadratic Equation Solver

::::{figure} images/quadratic-equation-solver.gif
:name: fig:quadratic
:alt: The quadratic equation solver
:align: center

Illustration of the quadratic equation solver

::::

::::{prf:proposition} quadratic equation
:label: quadratic

For the quadratic equation

$$
ax^2+bx+c=0,
$$ (eq:quadratic)
the *roots* (solutions for $x$) are give by

$$
\frac{-b-\sqrt{b^2-4ac}}{2a},\frac{-b+\sqrt{b^2-4ac}}{2a}.
$$ (eq:quadratic_roots)

::::

::::{exercise} roots
:label: ex:roots

Assign to `root1` and `root2` the values of the first and second roots in [](#eq:quadratic_roots) above respectively.

:::{caution}

Return the roots in the correct order and pay attention to the associativity of the arithmetic operators.

:::

::::

<!-- grade_id: get_roots, points: 0, grade: false, solution: true -->
```python
def get_roots(a, b, c):
    # YOUR CODE HERE
    raise NotImplementedError()
    return root1, root2
```

To test your code:

<!-- grade_id: test-get_roots, points: 1, grade: true, solution: false -->
```python
# tests
assert np.isclose(get_roots(1, 1, 0), (-1.0, 0.0)).all()
assert np.isclose(get_roots(1, 2, 1), (-1.0, -1.0)).all()
assert np.isclose(get_roots(2, 2, 1), (-0.5 - 0.5j, -0.5 + 0.5j)).all()
```

<!-- grade_id: htest-get_roots, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

To run the calculator illustrated in [](#fig:quadratic):

```python
# quadratic equations solver
@interact(a=(-10,10,1),b=(-10,10,1),c=(-10,10,1))
def quadratic_equation_solver(a=1,b=2,c=1):
    print('Roots: {}, {}'.format(*get_roots(a,b,c)))
```

## Number Conversion

### Byte-to-Decimal Calculator

::::{figure} images/byte-to-decimal.gif
:name: fig:byte-to-decimal
:alt: The byte-to-decimal calculator
:align: center

Illustration of the byte-to-decimal calculator

::::

Denote a binary number stored as a byte ($8$-bit) as

$$ 
b_7\circ b_6\circ b_5\circ b_4\circ b_3\circ b_2\circ b_1\circ b_0, 
$$
where $\circ$ concatenates $b_i$'s together into a binary string.

The binary string can be converted to a decimal number by the formula

$$ 
b_7\cdot 2^7 + b_6\cdot 2^6 + b_5\cdot 2^5 + b_4\cdot 2^4 + b_3\cdot 2^3 + b_2\cdot 2^2 + b_1\cdot 2^1 + b_0\cdot 2^0. 
$$

E.g., the binary string `'11111111'` is the largest integer represented by a byte:

$$
2^7+2^6+2^5+2^4+2^3+2^2+2^1+2^0=255=2^8-1.
$$

::::{exercise}

Assign to `decimal` the *integer* value represented by the binary sequence `b7,b6,b5,b4,b3,b2,b1,b0` of *characters* `'0'` or `'1'`.

::::

<!-- grade_id: byte_to_decimal, points: 0, grade: false, solution: true -->
```python
def byte_to_decimal(b7, b6, b5, b4, b3, b2, b1, b0):
    """
    Parameters
    ----------
    b7, ..., b0: single characters either '0' or '1'.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return decimal
```

To test your code:

<!-- grade_id: test-byte_to_decimal, points: 1, grade: true, solution: false -->
```python
# tests
def test_byte_to_decimal(decimal, b7, b6, b5, b4, b3, b2, b1, b0):
    decimal_ = byte_to_decimal(b7, b6, b5, b4, b3, b2, b1, b0)
    assert decimal == decimal_ and isinstance(decimal_, int)


test_byte_to_decimal(38, "0", "0", "1", "0", "0", "1", "1", "0")
test_byte_to_decimal(20, "0", "0", "0", "1", "0", "1", "0", "0")
test_byte_to_decimal(22, "0", "0", "0", "1", "0", "1", "1", "0")
test_byte_to_decimal(146, '1', '0', '0', '1', '0', '0', '1', '0')
test_byte_to_decimal(128, '1', '0', '0', '0', '0', '0', '0', '0')
test_byte_to_decimal(71, '0', '1', '0', '0', '0', '1', '1', '1')
```

<!-- grade_id: htest-byte_to_decimal, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

To run the calculator illustrate in [](#fig:byte-to-decimal):

```python
# byte-to-decimal calculator
bit = ['0', '1']


@interact(b7=bit, b6=bit, b5=bit, b4=bit, b3=bit, b2=bit, b1=bit, b0=bit)
def convert_byte_to_decimal(b7, b6, b5, b4, b3, b2, b1, b0):
    print('decimal:', byte_to_decimal(b7, b6, b5, b4, b3, b2, b1, b0))
```

### Decimal-to-Byte Calculator

::::{figure} images/decimal-to-byte.gif
:name: fig:decimal-to-byte
:alt: The decimal-to-byte calculator
:align: center

Illustration of the decimal-to-byte calculator

::::

::::{exercise}
:label: ex:decimal-to-byte

Assign to `byte` a *string of 8 bits* that represents the value of `decimal`, a non-negative decimal integer from $0$ to $2^8-1=255$.  

:::{important}

Use the operators `//` and `%`, but not the function `int`.
:::

::::

<!-- grade_id: decimal_to_byte, points: 0, grade: false, solution: true -->
```python
def decimal_to_byte(decimal):
    # YOUR CODE HERE
    raise NotImplementedError()
    return byte
```

To test your code:

<!-- grade_id: test-decimal_to_byte, points: 1, grade: true, solution: false -->
```python
# tests
def test_decimal_to_byte(byte, decimal):
    byte_ = decimal_to_byte(decimal)
    assert byte == byte_ and isinstance(byte, str) and len(byte) == 8


test_decimal_to_byte("01100111", 103)
test_decimal_to_byte("00000011", 3)
test_decimal_to_byte("00011100", 28)
test_decimal_to_byte('11011111', 223)
test_decimal_to_byte('00000100', 4)
test_decimal_to_byte('10011001', 153)


def test_decimal_to_byte(byte,decimal):
    byte_ = decimal_to_byte(decimal)
    correct = byte == byte_ and isinstance(byte, str) and len(byte) == 8
    if not correct:
        print(
            f'{decimal} should be represented as the byte {byte}, not {byte_}.'
        )
    assert correct


test_decimal_to_byte('01100111', 103)
test_decimal_to_byte('00000011', 3)
test_decimal_to_byte('00011100', 28)
```

<!-- grade_id: htest-decimal_to_byte, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

To run the calculator illustrate in [](#fig:decimal-to-byte):

```python
# decimal-to-byte calculator
@interact(decimal=(0, 255, 1))
def convert_decimal_to_byte(decimal=0):
    print("byte:", decimal_to_byte(decimal))
```

## Symbolic calculators

Can we do complicated arithmetics with python. What about Calculus?

$$
\int \tan(x)\, dx = \color{red}{?}
$$

Try [SymPy Gamma](https://gamma.sympy.org/input/?i=integrate%28tan%28x%29%29) or [SymPy Beta](https://sympy-beta.vercel.app/input/integrate%28tan%28x%29%29).

::::{tip}

- Take a look at the different panels to learn about the solution: `Steps`, `Plot`, and `Derivative`.
- Try different [random examples](https://gamma.sympy.org/).

::::

**How does SymPy Gamma work?**

- [SymPy Gamma](https://gamma.sympy.org/) is a web application running [SymPy](https://docs.sympy.org/latest/index.html), which is a python library for symbolic computation.
- [SymPy Beta](https://github.com/eagleoflqj/sympy_beta) is a fork of SymPy Gamma that can run totally in the browser.

**How to use SymPy?**

The following line in the initialization cell imports the library:

```python
import sympy as sp
```

We need to define a symbolic variable and assign it to a python variable.

```python
x = sp.symbols("x")
x
```

The SymPy expression for $\tan(x)$ is:

```python
f = sp.tan(x)
f
```

To compute the integration:

$$
\int f(x) dx
$$

```python
g = sp.integrate(f)
g
```

To compute the derivative:

$$
\frac{d}{dx}g(x)
$$

```python
diff_g = sp.diff(g)
diff_g
```

The answer can be simplified as expected:

```python
diff_g.simplify()
```

To plot the functions:

```python
p = sp.plot(f, g, (x, -2 * sp.pi / 5, 2 * sp.pi / 5), ylabel="y", legend=True)
```

::::{exercise} sympy
:label: ex:sympy

Using SymPy expressions

- assign to `x` a SymPy variable named `"x"`,
- assign to `f` the expression $\frac{1}{\sqrt{1 - x^2}}$ in terms of `x`,
- assign to `g` the result of $\int \frac{1}{\sqrt{1 - x^2}} dx$, and
- *optionally*, plot `f` and `g` for $x\in \left[-\frac45, \frac45\right]$.

:::{caution}

Use `sp.sqrt` or `**(sp.S(1)/2)` for square root instead of `**0.5`. See [SymPy gotchas](https://docs.sympy.org/latest/gotchas.html).
:::
::::

<!-- grade_id: sympy, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

The following test should plot your expression `f` in SymPy.

<!-- grade_id: test-sympy, points: 1, grade: true, solution: false -->
```python
# tests
assert sp.simplify(f.subs(x, 0) - 1) == 0
assert sp.simplify(f.subs(x, sp.S(1)/2) - sp.sqrt(3)*2/3) == 0
```

<!-- grade_id: htest-sympy, points: 1, grade: true, solution: false -->
```python
# hidden tests
```